**주의사항**
 - 제출 전 커널을 재시작하고 **run all cells**를 눌러 모든 코드가 정상적으로 동작하는지 확인하세요.
 - 이름을 기입하고 Shift + Enter 를 누르세요.
 - 파일을 저장했는지 꼭 확인하세요.
 - 함수에서 `YOUR CODE HERE`라고 쓰여 있는 부분을 지우고 함수를 작성하세요.
 - 문제 아래에 테스트 코드를 실행해서 오류가 없는 지 확인해 보세요.
 - 테스트 코드를 통과했다고 해서 반드시 정답인 것은 아닙니다.

```python
이름 = ""
```

---

### 문제 1.

문자열에서 콤마로 구분된 단어를 추출하고 공백을 제거한 후, 다시 단어를 공백으로 연결하는 함수 `process_string`를 작성하세요.

- 함수는 하나의 매개변수를 받습니다: `input_str` (처리할 문자열).
- 함수는 문자열을 콤마 `,` 로 분리하고, 각 단어의 앞뒤 공백을 제거해야 합니다.
- 마지막으로, 단어들을 하나의 문자열로 `공백`으로 이어붙여 반환합니다.
- `,`가 있었던 자리가 공백으로 처리되어 있을 수 있기 때문에 문자열 양끝에 공백이 있다면 다시 한 번 처리해 줘야 합니다.

예를 들어:

- `input_str = "  apple, banana , cherry  "`의 경우 출력은 `"apple banana cherry"`이어야 합니다.
- `input_str = "  hello, world , "`의 경우 출력은 `"hello world"`이어야 합니다.


<!-- grade_id: a1a, points: 0, grade: false, solution: true -->
```python
def process_string(input_str: str) -> str:
    ### BEGIN SOLUTION
    s = [i.strip(" ") for i in input_str.split(",")]
    return " ".join(s).strip(' ')
    ### END SOLUTION
```

<!-- grade_id: a1t, points: 3, grade: true, solution: false -->
```python
assert process_string("  apple, banana , cherry  ") == "apple banana cherry"
assert process_string("dog,  cat ,bird") == "dog cat bird"
assert process_string("  hello, world , ") == "hello world"

### BEGIN HIDDEN TESTS
assert process_string("  hello, world , guys") == "hello world guys"
### END HIDDEN TESTS
```

### 문제 2.

문자열을 쉼표로 분리하여 리스트로 만든 후, 리스트의 두 번째부터 뒤에서 두 번째 요소까지 선택하여 정렬한 새로운 리스트를 생성하는 `sort_elements` 함수를 작성하세요.


- 함수는 하나의 매개변수를 받습니다: `input_str` (처리할 문자열).
- 함수는 먼저 문자열을 쉼표 ,로 분리하여 리스트를 생성합니다.
- 그 다음, 리스트의 두 번째부터 뒤에서 두 번째 요소까지를 선택합니다.
- 마지막으로, 선택한 요소들을 알파벳 순서로 정렬하여 새로운 리스트를 생성하여 반환합니다.

예를 들어:

`input_str = "g,f,e,d,c,b,a"`의 경우 출력은 `['b', 'c', 'd', 'e', 'f']`이어야 합니다.


<!-- grade_id: a2a, points: 0, grade: false, solution: true -->
```python
def sort_elements(input_str: str) -> list:
    ### BEGIN SOLUTION
    elements = input_str.split(',')
    selected_elements = elements[1:-1]
    sorted_elements = sorted(selected_elements)
    return sorted_elements
    ### END SOLUTION
```

<!-- grade_id: a2t, points: 3, grade: true, solution: false -->
```python
assert sort_elements("g,f,e,d,c,b,a") == ['b', 'c', 'd', 'e', 'f']
assert sort_elements("m,l,k,j,i,h,g") == ['h', 'i', 'j', 'k', 'l']
### BEGIN HIDDEN TESTS
assert sort_elements("z,y,x,w,v,u,t") == ['u', 'v', 'w', 'x', 'y']
### END HIDDEN TESTS
```

### 문제 3.

문자열 `s`와 패턴 `p`가 주어질 때, 패턴 `p`가 문자열 `s`에 몇 번 나타나는지 카운트하는 함수 `count_pattern`를 작성하세요.

- 함수는 두 개의 매개변수를 받습니다: `s` (문자열), `p` (찾을 패턴의 문자열)
- 함수는 문자열 `s`에서 패턴 `p`가 나타나는 횟수를 반환합니다.

예를 들어:

- `s = "abcabcabc", p = "abc"`의 경우 출력은 `3`이어야 합니다.


<!-- grade_id: a3a, points: 0, grade: false, solution: true -->
```python
def count_pattern(s: str, p: str) -> int:
    ### BEGIN SOLUTION
    count = 0
    for i in range(len(s) - len(p) + 1):
        if s[i:i + len(p)] == p:
            count += 1
    return count
    ### END SOLUTION
```

<!-- grade_id: a3t, points: 3, grade: true, solution: false -->
```python
assert count_pattern("abcabcabc", "abc") == 3
assert count_pattern("aaaaaa", "aa") == 5
### BEGIN HIDDEN TESTS
assert count_pattern("xyzxyz", "xyz") == 2
### END HIDDEN TESTS
```

### 문제 4.

우리 건물은 층마다 센서가 여러 개 달려 있습니다. 우리는 이 센서들에서 값을 가져오려고 하고 이것을 2차원 리스트로 만들었습니다.

각 행은 건물의 층을, 각 열은 해당 층의 센서 값을 나타냅니다.

각 층에서 첫 번째, 중간, 마지막 센서의 데이터만 추출하는 extract_sensor_data 함수를 작성하세요.

- 함수는 하나의 매개변수를 받습니다: `sensor_data` (건물의 각 층에 있는 센서 데이터를 담은 2D 리스트).
- 함수는 각 층에서 첫 번째, 중간(또는 중간 두 개의 평균), 마지막 센서의 데이터만 추출하여 새로운 이차원 리스트로 반환합니다.
- 만약 센서의 개수가 짝수개라면 중간값은 가운데 두 센서 값의 평균을 사용합니다. 홀수개라면 정확히 중간의 센서 값을 사용합니다.
예를 들어:

`sensor_data = [[1, 2, 3, 4], [5, 6, 7, 8], [9, 10, 11, 12]]`의 경우 출력은 `[[1, 2.5, 4], [5, 6.5, 8], [9, 10.5, 12]]`이어야 합니다.
`sensor_data = [1, 3, 5], [2, 4, 6], [7, 9, 11]`의 출력은 `[[1, 3, 5], [2, 4, 6], [7, 9, 11]]`이어야 합니다.



<!-- grade_id: a4a, points: 0, grade: false, solution: true -->
```python
def extract_sensor_data(sensor_data: list) -> list:
    ### BEGIN SOLUTION
    extracted_data = []
    for row in sensor_data:
        first = row[0]
        if len(row) % 2 == 1:
            middle = row[len(row)//2]
        else:
            middle = (row[len(row)//2 - 1] + row[len(row)//2]) / 2
        last = row[-1]
        extracted_data.append([first, middle, last])
    return extracted_data
    ### END SOLUTION
```

<!-- grade_id: a4t, points: 3, grade: true, solution: false -->
```python
assert extract_sensor_data([[1, 2, 3, 4], [5, 6, 7, 8], [9, 10, 11, 12]]) == [[1, 2.5, 4], [5, 6.5, 8], [9, 10.5, 12]]
assert extract_sensor_data([[1, 3, 5], [2, 4, 6], [7, 9, 11]]) == [[1, 3, 5], [2, 4, 6], [7, 9, 11]]
### BEGIN HIDDEN TESTS
assert extract_sensor_data([[10, 20, 30], [40, 50, 60], [70, 80, 90]]) == [[10, 20, 30], [40, 50, 60], [70, 80, 90]]
assert extract_sensor_data([[10, 20, 30, 40], [50, 60, 70, 80], [90, 100, 110, 120]]) == [[10, 25, 40], [50, 65, 80], [90, 105, 120]]
### END HIDDEN TESTS
```

### 문제 5.

소설 '반지의 제왕'에서, 프로도와 샘은 반지를 모르도르로 가져가기 위해 여정을 떠납니다. 그들은 각자 다른 지식과 능력을 가지고 있습니다. 어떤 능력은 공통적이지만, 일부는 서로 다릅니다. 이 두 캐릭터가 가지고 있는 능력 중에서, 둘 중 한 명만이 가지고 있는 능력을 찾아야 악의 무리에 대항할 수 있습니다.

함수 `find_unique_abilities`를 작성하세요.

- 함수는 두 매개변수를 받습니다:
  1. `frodo_abilities` (프로도가 가지고 있는 능력을 문자열로 저장하는 set)
  2. `sam_abilities` (샘이 가지고 있는 능력을 문자열로 저장하는 set)

- 함수는 두 캐릭터 중 한 명만이 가지고 있는 능력의 목록을 `set` 형태로 반환합니다.

예를 들어:

- `frodo_abilities = {"Ring Bearing", "Stealth", "Courage"}`
- `sam_abilities = {"Cooking", "Courage", "Loyalty"}`

함수의 출력은 `{"Ring Bearing", "Stealth", "Cooking", "Loyalty"}`이어야 합니다.



<!-- grade_id: a5a, points: 0, grade: false, solution: true -->
```python
def find_unique_abilities(frodo_abilities: set, sam_abilities: set) -> set:
    ### BEGIN SOLUTION
    return frodo_abilities.symmetric_difference(sam_abilities)
    ### END SOLUTION
```

<!-- grade_id: a5t, points: 3, grade: true, solution: false -->
```python
assert find_unique_abilities({"Ring Bearing", "Stealth", "Courage"}, {"Cooking", "Courage", "Loyalty"}) == {"Ring Bearing", "Stealth", "Cooking", "Loyalty"}
### BEGIN HIDDEN TESTS
assert find_unique_abilities({"Swordsmanship", "Archery"}, {"Cooking", "Healing"}) == {"Swordsmanship", "Archery", "Cooking", "Healing"}
### END HIDDEN TESTS
```

### 문제 6.

과학자인 당신은 대기 질 모니터링을 위해 공기질 센서를 설치하였습니다. 이 센서는 일정 간격으로 공기 중 미세먼지 농도를 측정하여 리스트로 저장합니다.

함수 `filter_readings`를 작성하세요.

- 함수는 하나의 매개변수를 받습니다: `readings` (미세먼지 농도의 리스트)
- 함수는 미세먼지 농도가 `35` 이상인 측정치만 추출하여 새로운 리스트를 생성합니다.
- 함수는 새로운 리스트를 반환합니다.

예를 들어:

- `readings = [25, 40, 35, 18, 50]`의 경우 출력은 `[40, 35, 50]`이어야 합니다.


<!-- grade_id: a6a, points: 0, grade: false, solution: true -->
```python
def filter_readings(readings: list) -> list:
    ### BEGIN SOLUTION
    filtered_readings = [reading for reading in readings if reading >= 35]
    return filtered_readings
    ### END SOLUTION
```

<!-- grade_id: a6t, points: 3, grade: true, solution: false -->
```python
assert filter_readings([25, 40, 35, 18, 50]) == [40, 35, 50]
assert filter_readings([10, 20, 30]) == []
### BEGIN HIDDEN TESTS
assert filter_readings([36, 40, 22, 37, 18]) == [36, 40, 37]
### END HIDDEN TESTS
```

### 문제 7.

온라인 쇼핑몰에서 주문 내역이 딕셔너리 형태로 저장되어 있습니다.

함수 `process_orders`를 작성하세요.

- 함수는 하나의 매개변수를 받습니다: `orders` (고객의 주문 내역을 저장하는 딕셔너리)
- 함수는 각 고객의 주문 금액을 계산하여 새로운 딕셔너리를 생성합니다. 주문 금액은 상품의 가격과 수량을 곱한 값입니다.
- 딕셔너리의 키는 고객의 이메일, 값은 주문 금액입니다.
- 만약 주문 내역에 `price`값이 없으면, `price`를 0으로 가정합니다.

예를 들어:

- `orders = {"john@email.com": {"item": "book", "quantity": 1, "price": 10}, "jane@email.com": {"item": "pen", "quantity": 3}}`의 경우 출력은 `{"john@email.com": 10, "jane@email.com": 0}`이어야 합니다.


<!-- grade_id: a7a, points: 0, grade: false, solution: true -->
```python
def process_orders(orders: dict) -> dict:
    ### BEGIN SOLUTION
    processed_orders = {}
    for email, order in orders.items():
        price = order.get("price", 0)
        quantity = order.get("quantity", 0)
        total_price = price * quantity
        processed_orders[email] = total_price
    return processed_orders
    ### END SOLUTION
```

<!-- grade_id: a7t, points: 3, grade: true, solution: false -->
```python
assert process_orders({"john@email.com": {"item": "book", "quantity": 1, "price": 10}, "jane@email.com": {"item": "pen", "quantity": 3}}) == {"john@email.com": 10, "jane@email.com": 0}
assert process_orders({"mary@email.com": {"item": "laptop", "quantity": 1, "price": 1000}}) == {"mary@email.com": 1000}
### BEGIN HIDDEN TESTS
assert process_orders({"dave@email.com": {"item": "phone", "quantity": 2, "price": 500}, "ella@email.com": {"item": "headphones", "quantity": 1}}) == {"dave@email.com": 1000, "ella@email.com": 0}
### END HIDDEN TESTS
```

### 문제 8.

당신은 포켓몬 트레이너입니다. 내가 잡은 포켓몬을 도감에 넣고 관리하고자 합니다. 

함수 `my_pokemon`를 작성하세요.

- 함수는 두 개의 매개변수를 받습니다: `name` (포켓몬의 이름), `**kwargs` (키워드 파라미터로 받는 추가 정보)
- 함수는 `name`과 추가로 입력받은 정보를 합쳐서 하나의 딕셔너리로 만듭니다.
- 함수는 이 딕셔너리를 반환합니다.

예를 들어:

- `my_pokemon("Pikachu", level=3, color="yellow")`의 경우 출력은 `{'name': 'Pikachu', 'level': 30, 'color': 'yellow'}`이어야 합니다.
-  `my_pokemon("Ggobugi", level=1, color="blue", type="water")`의 경우 출력은 `{'name': 'Ggobugi', 'level': 1, 'color': 'blue', 'type':'water'}`이어야 합니다.


<!-- grade_id: a8a, points: 0, grade: false, solution: true -->
```python
def my_pokemon(name: str, **kwargs) -> dict:
    ### BEGIN SOLUTION
    info = {"name": name}
    info.update(kwargs)
    return info
    ### END SOLUTION
```

<!-- grade_id: a8t, points: 3, grade: true, solution: false -->
```python
assert my_pokemon("Pikachu", level=3, color="yellow") == {'name': 'Pikachu', 'level': 3, 'color': 'yellow'}
assert my_pokemon("Ggobugi", level=1, color="blue", type="water") == {'name': 'Ggobugi', 'level': 1, 'color': 'blue', 'type':'water'}
### BEGIN HIDDEN TESTS
assert my_pokemon("최동혁", level=12, color="yellow") == {'name': '최동혁', 'level': 12, 'color': 'yellow'}
### END HIDDEN TESTS
```

### 문제 9.

비행기의 출발 시간과 도착 시간의 차이를 분 단위로 계산해고 싶습니다.

함수 `calculate_flight_duration`를 작성하세요.

- 함수는 두 개의 매개변수를 받습니다: `departure_time` (출발 시간을 나타내는 datetime 객체), `arrival_time` (도착 시간을 나타내는 datetime 객체)
- 함수는 출발 시간과 도착 시간의 차이를 분 단위로 계산하여 반환합니다.

예를 들어:

- `departure_time = datetime(2023, 1, 1, 10, 0, 0)`, `arrival_time = datetime(2023, 1, 1, 12, 30, 0)`의 경우 출력은 `150`이어야 합니다.


<!-- grade_id: a9a, points: 0, grade: false, solution: true -->
```python
from datetime import datetime

def calculate_flight_duration(departure_time: datetime, arrival_time: datetime) -> int:
    ### BEGIN SOLUTION
    delta = arrival_time - departure_time
    return delta.seconds // 60
    ### END SOLUTION
```

<!-- grade_id: a9t, points: 3, grade: true, solution: false -->
```python
assert calculate_flight_duration(datetime(2023, 1, 1, 10, 0, 0), datetime(2023, 1, 1, 12, 30, 0)) == 150
assert calculate_flight_duration(datetime(2023, 1, 1, 9, 0, 0), datetime(2023, 1, 1, 9, 45, 0)) == 45
### BEGIN HIDDEN TESTS
assert calculate_flight_duration(datetime(2023, 1, 2, 11, 0, 0), datetime(2023, 1, 2, 12, 15, 0)) == 75
### END HIDDEN TESTS
```

### 문제 10.

당신은 딕셔너리에서 key값과 value값을 반대로 입력했다는 사실을 깨달았습니다. 이를 수정하는 함수 `reverse_key_value`를 작성하세요.

- 함수는 하나의 매개변수를 받습니다: `original_dict` (원본 키와 값을 저장하는 딕셔너리).
- 함수는 원본 딕셔너리의 키와 값이 서로 바뀐 새로운 딕셔너리를 반환합니다.

예를 들어:

- `original_dict = {"apple": "fruit", "lettuce": "vegetable"}`의 경우 출력은 `{"fruit": "apple", "vegetable": "lettuce"}`이어야 합니다.

**주의**: 이 문제에서는 원본 딕셔너리의 모든 값이 유니크하다고 가정합니다.



<!-- grade_id: a10a, points: 0, grade: false, solution: true -->
```python
def reverse_key_value(original_dict: dict) -> dict:
    ### BEGIN SOLUTION
    reversed_dict = {}
    for key, value in original_dict.items():
        reversed_dict[value] = key
    return reversed_dict
    ### END SOLUTION
```

<!-- grade_id: a10t, points: 3, grade: true, solution: false -->
```python
assert reverse_key_value({"apple": "fruit", "lettuce": "vegetable"}) == {"fruit": "apple", "vegetable": "lettuce"}
assert reverse_key_value({"cat": "animal", "rose": "flower"}) == {"animal": "cat", "flower": "rose"}
### BEGIN HIDDEN TESTS
assert reverse_key_value({"book": "read", "pencil": "write"}) == {"read": "book", "write": "pencil"}
### END HIDDEN TESTS
```