**주의사항**
 - 제출 전 커널을 재시작하고 **run all cells**를 눌러 모든 코드가 정상적으로 동작하는지 확인하세요.
 - 이름을 기입하고 Shift + Enter 를 누르세요.
 - 파일을 저장했는지 꼭 확인하세요.
 - 함수에서 `YOUR CODE HERE`라고 쓰여 있는 부분을 지우고 함수를 작성하세요.
 - 문제 아래에 테스트 코드를 실행해서 오류가 없는 지 확인해 보세요.
 - 테스트 코드를 통과했다고 해서 반드시 정답인 것은 아닙니다.

```python
이름 = ""
```

---

### 문제 1.

당신은 주중에는 6시 30분에 일어나고 싶고, 주말에는 자유롭게 잠을 자고 싶습니다.

함수 `set_alarm`를 작성하세요.

- 함수는 하나의 매개변수를 받습니다: `weekday` (Boolean 값, 주중이면 `True`, 주말이면 `False`)
- 함수는 주중일 경우 "wake up at 6:30"을, 주말일 경우 "sleep in"을 반환합니다.

예를 들어:

- `weekday = True`일 경우, 출력은 "wake up at 6:30"이어야 합니다.
- `weekday = False`일 경우, 출력은 "sleep in"이어야 합니다.

<!-- grade_id: b1a, points: 0, grade: false, solution: true -->
```python
def set_alarm(weekday: bool) -> str:
    ### BEGIN SOLUTION
    if weekday:
        return "wake up at 6:30"
    else:
        return "sleep in"
    ### END SOLUTION
```

<!-- grade_id: b1t, points: 1, grade: true, solution: false -->
```python
assert set_alarm(True) == "wake up at 6:30"
assert set_alarm(False) == "sleep in"
### BEGIN HIDDEN TESTS
assert set_alarm(True) == "wake up at 6:30"
assert set_alarm(False) == "sleep in"
### END HIDDEN TESTS
```

### 문제 2.

함수 `get_length`를 작성하세요.

주어진 리스트 `lst`의 길이를 반환하세요.

예를 들어:

- `lst = [10, 20, 30, 40, 50]`의 경우 출력은 `5`이어야 합니다.

- `lst = ["안녕", "하세요", "파이썬"]`의 경우 출력은 `3`이어야 합니다.


<!-- grade_id: b2a, points: 0, grade: false, solution: true -->
```python
def get_length(lst: list) -> int:
    ### BEGIN SOLUTION
    return len(lst)
    ### END SOLUTION
```

<!-- grade_id: b2t, points: 1, grade: true, solution: false -->
```python
assert get_length([10, 20, 30, 40, 50]) == 5
assert get_length(["안녕", "하세요", "파이썬"]) == 3
assert get_length([True, False, True, True, False]) == 5
assert get_length(["데이터", "베이스", "시스템", "프로그래밍"]) == 4
### BEGIN HIDDEN TESTS
assert get_length(["이것", "은", "테스트", "입니다"]) == 4
assert get_length([1, 2, 3, 4, 5, 6, 7, 8, 9, 10]) == 10
### END HIDDEN TESTS 
```

### 문제 3.

당신은 자주 가는 카페와 그 카페의 메뉴에 대한 정보를 문자열로 표현하려고 합니다.

함수 `favorite_cafe`를 작성하세요.

- 함수는 두 개의 매개변수를 받습니다: `brand` (가장 좋아하는 카페 브랜드의 이름, 문자열)와 `menu` (그 카페에서 가장 좋아하는 메뉴, 문자열)
- 함수는 f-string을 이용하여 `"제가 제일 좋아하는 카페는 {brand}입니다. 그 중에서도 {menu}가 가장 좋아요."` 형태의 문자열을 반환합니다.

예를 들어:

- `brand = "Starbucks", menu = "Americano"`일 경우, 출력은 "제가 제일 좋아하는 카페는 Starbucks입니다. 그 중에서도 Americano가 가장 좋아요."이어야 합니다.



<!-- grade_id: b3a, points: 0, grade: false, solution: true -->
```python
def favorite_cafe(brand: str, menu: str) -> str:
    ### BEGIN SOLUTION
    return f'제가 제일 좋아하는 카페는 {brand}입니다. 그 중에서도 {menu}가 가장 좋아요.'
    ### END SOLUTION
```

<!-- grade_id: b3t, points: 1, grade: true, solution: false -->
```python
assert favorite_cafe("Starbucks", "Americano") == "제가 제일 좋아하는 카페는 Starbucks입니다. 그 중에서도 Americano가 가장 좋아요."
assert favorite_cafe("Ediya", "Iced Latte") == "제가 제일 좋아하는 카페는 Ediya입니다. 그 중에서도 Iced Latte가 가장 좋아요."
### BEGIN HIDDEN TESTS
assert favorite_cafe("Hollys", "Cappuccino") == "제가 제일 좋아하는 카페는 Hollys입니다. 그 중에서도 Cappuccino가 가장 좋아요."
### END HIDDEN TESTS
```

### 문제 4.

당신은 두 개의 문장을 합쳐서 하나의 완전한 문장을 만들려고 합니다.

함수 `concatenate_strings`를 작성하세요.

- 함수는 두 개의 매개변수를 받습니다: `first` (첫 번째 문장, 문자열)와 `second` (두 번째 문장, 문자열)
- 함수는 두 개의 문자열을 연결하여 반환합니다.

예를 들어:

- `first = "국경의 긴 터널을 빠져나오자, ", second = "설국이었다."`일 경우, 출력은 "국경의 긴 터널을 빠져나오자, 설국이었다."이어야 합니다.



<!-- grade_id: b4a, points: 0, grade: false, solution: true -->
```python
def concatenate_strings(first: str, second: str) -> str:
    ### BEGIN SOLUTION
    return first + second
    ### END SOLUTION
```

<!-- grade_id: b4t, points: 1, grade: true, solution: false -->
```python
assert concatenate_strings("국경의 긴 터널을 빠져나오자, ", "설국이었다.") == "국경의 긴 터널을 빠져나오자, 설국이었다."
assert concatenate_strings("오늘은 ", "월요일입니다.") == "오늘은 월요일입니다."
### BEGIN HIDDEN TESTS
assert concatenate_strings("나는 ", "학생입니다.") == "나는 학생입니다."
### END HIDDEN TESTS
```

### 문제 5.

아래 initial_list에 저장된 값을 변경하고자 합니다. `"Madison"`은 진급 시험을 통과했습니다. initial_list 명단의 Grade Level을 7학년이 아닌 8학년으로 바꿔야 합니다. 변경사항을 반영하기 위해 `[][]`를 사용하여 목록을 수정하세요. **양수 인덱스**를 사용하하세요. 또, `"Kenny"`는 별명인 `"Ken"`으로 불리기를 좋아합니다. 리스트를 **음수 인덱스**만 사용하여 변경해 보세요.

**변경 전**
```
initial_list = [["Kenny", "America", 9],
                ["Tanya", "Ukrainian", 9],
                ["Madison", "Indian", 7]]

```
**변경 후**
```
initial_list = [["Ken", "America", 9],
                ["Tanya", "Ukrainian", 9],
                ["Madison", "Indian", 8]]

```


<!-- grade_id: b5a, points: 0, grade: false, solution: true -->
```python
def update_incoming_class(incoming_class):
    ### BEGIN SOLUTION
    incoming_class[2][2] = 8  # Using positive index to update Madison's grade
    incoming_class[-3][0] = "Ken"  # Using negative index to update Kenny's name
    return incoming_class
    ### END SOLUTION
```

<!-- grade_id: b5t, points: 1, grade: true, solution: false -->
```python
initial_list = [["Kenny", "America", 9],
                ["Tanya", "Ukrainian", 9],
                ["Madison", "Indian", 7]]

assert update_incoming_class(initial_list) == [["Ken", "America", 9],
                                               ["Tanya", "Ukrainian", 9],
                                               ["Madison", "Indian", 8]]

### BEGIN HIDDEN TESTS
assert update_incoming_class(initial_list)[0][0] == "Ken"
assert update_incoming_class(initial_list)[2][2] == 8
### END HIDDEN TESTS
```

### 문제 6.

온도 변환 함수를 작성하세요. 화씨 온도를 섭씨 온도로 변환하는 기능을 구현합니다. 

함수 `f_to_c`를 작성하세요.

- 함수는 하나의 매개변수를 받습니다: `f_temp` (화씨 온도, 실수형)
- 함수는 화씨 온도를 섭씨로 변환한 후, 그 값을 반환합니다. 섭씨 온도는 `c_temp`로 표시하며, 다음의 수식을 사용하세요:
  
  $[
  \text{온도 (C)} = \left( \text{온도 (F)} - 32 \right) \times \frac{5}{9}
  ]$
  
예를 들어:

- `f_temp = 32.0`의 경우, 출력은 `0.0`이어야 합니다.
- `f_temp = 212.0`의 경우, 출력은 `100.0`이어야 합니다.

<!-- grade_id: b6a, points: 0, grade: false, solution: true -->
```python
def f_to_c(f_temp: float) -> float:
    ### BEGIN SOLUTION
    c_temp = (f_temp - 32) * 5 / 9
    return c_temp
    ### END SOLUTION
```

<!-- grade_id: b6t, points: 1, grade: true, solution: false -->
```python
assert f_to_c(32.0) == 0.0
assert f_to_c(212.0) == 100.0
### BEGIN HIDDEN TESTS
assert f_to_c(98.6) == 37.0
### END HIDDEN TESTS
```

### 문제 7.

온라인 쇼핑몰에서 고객이 장바구니에 상품을 추가하는 기능을 구현해야 합니다.

함수 `add_to_cart`를 작성하세요.

- 함수는 두 개의 매개변수를 받습니다: `cart` (장바구니에 담긴 상품 리스트)와 `item` (장바구니에 추가할 상품, 문자열)
- 함수는 `item`을 `cart` 리스트의 마지막에 추가한 후, 그 결과 리스트를 반환합니다.

예를 들어:

- `cart = ["사과", "바나나", "딸기"]`, `item = "포도"`의 경우 출력은 `["사과", "바나나", "딸기", "포도"]`이어야 합니다.
- `cart = ["치즈", "버터"]`, `item = "우유"`의 경우 출력은 `["치즈", "버터", "우유"]`이어야 합니다.


<!-- grade_id: b7a, points: 0, grade: false, solution: true -->
```python
def add_to_cart(cart: list[str], item: str) -> list[str]:
    ### BEGIN SOLUTION
    cart.append(item)
    return cart
    ### END SOLUTION
```

<!-- grade_id: b7t, points: 1, grade: true, solution: false -->
```python
assert add_to_cart(["사과", "바나나", "딸기"], "포도") == ["사과", "바나나", "딸기", "포도"]
assert add_to_cart(["치즈", "버터"], "우유") == ["치즈", "버터", "우유"]
### BEGIN HIDDEN TESTS
assert add_to_cart([], "빵") == ["빵"]
### END HIDDEN TESTS
```

### 문제 8.

파이썬의 `if`문을 사용하여 주어진 숫자가 짝수인지 홀수인지 판별하는 함수를 작성하세요.

함수 `even_or_odd`를 작성하세요.

이 함수는 정수 `num`을 매개변수로 받습니다. 만약 `num`이 짝수라면 "짝수"를 반환하고, 홀수라면 "홀수"를 반환하세요.

예를 들어:

- `num = 4`의 경우 출력은 "짝수"이어야 합니다.

- `num = 7`의 경우 출력은 "홀수"이어야 합니다.


<!-- grade_id: b8a, points: 0, grade: false, solution: true -->
```python
def even_or_odd(num: int) -> str:
    ### BEGIN SOLUTION
    if num % 2 == 0:
        return "짝수"
    else:
        return "홀수"
    ### END SOLUTION
```

<!-- grade_id: b8t, points: 1, grade: true, solution: false -->
```python
assert even_or_odd(4) == "짝수"
assert even_or_odd(7) == "홀수"
assert even_or_odd(10) == "짝수"
assert even_or_odd(13) == "홀수"
### BEGIN HIDDEN TESTS
assert even_or_odd(100) == "짝수"
assert even_or_odd(77) == "홀수"
### END HIDDEN TESTS 
```

### 문제 9.

어떤 놀이공원에서는 놀이기구를 이용하기 위한 키 제한이 있습니다. 여러분은 해당 놀이기구를 탈 수 있는지의 여부를 판단하는 프로그램을 작성해야 합니다.

함수 `can_ride`를 작성하세요.

- 함수는 두 개의 정수 `height` (탑승자의 키)와 `age` (탑승자의 나이)를 매개변수로 받습니다.
- 다음의 조건에 따라 문자열을 반환합니다:
  - 만약 `height`가 140cm 이상이고 `age`가 10세 이상이면 "탑승 가능"을 반환합니다.
  - 만약 `height`가 120cm 미만이거나 `age`가 5세 미만이면 "탑승 불가"를 반환합니다.
  - 위의 두 조건 모두 만족하지 않는 경우 "보호자 동반 시 탑승 가능"을 반환합니다.

예를 들어:

- `height = 145`, `age = 12`의 경우 출력은 "탑승 가능"이어야 합니다.
- `height = 115`, `age = 6`의 경우 출력은 "탑승 불가"이어야 합니다.
- `height = 130`, `age = 7`의 경우 출력은 "보호자 동반 시 탑승 가능"이어야 합니다.


<!-- grade_id: b9a, points: 0, grade: false, solution: true -->
```python
def can_ride(height: int, age: int) -> str:
    ### BEGIN SOLUTION
    if height >= 140 and age >= 10:
        return "탑승 가능"
    elif height < 120 or age < 5:
        return "탑승 불가"
    else:
        return "보호자 동반 시 탑승 가능"
    ### END SOLUTION
```

<!-- grade_id: b9t, points: 1, grade: true, solution: false -->
```python
assert can_ride(145, 12) == "탑승 가능"
assert can_ride(115, 6) == "탑승 불가"
assert can_ride(130, 7) == "보호자 동반 시 탑승 가능"
assert can_ride(138, 9) == "보호자 동반 시 탑승 가능"
### BEGIN HIDDEN TESTS
assert can_ride(142, 11) == "탑승 가능"
assert can_ride(118, 4) == "탑승 불가"
assert can_ride(124, 6) == "보호자 동반 시 탑승 가능"
### END HIDDEN TESTS 
```

### 문제 10.

당신은 노래 가사를 분석하는 프로그램을 작성하고 있습니다. 가사 내에서 특정 문자를 기준으로 문장을 나누려고 합니다.

함수 `split_lyrics`를 작성하세요.

- 함수는 두 개의 매개변수를 받습니다: `lyrics` (분석할 가사를 나타내는 문자열), `split_char` (가사를 나눌 기준 문자)
- 함수는 `split_char`을 기준으로 `lyrics`를 나눈 후, 그 결과 리스트를 반환합니다.

예를 들어:

- `lyrics = "나는 사과, 당근, 바나나를 좋아해"`, `split_char = ","`의 경우 출력은 `["나는 사과", " 당근", " 바나나를 좋아해"]`이어야 합니다.

- `lyrics = "반짝반짝 작은별 아름답게 비치네"` , `split_char = " "`의 경우 출력은 `["반짝반짝", "작은별", "아름답게", "비치네"]`이어야 합니다.


<!-- grade_id: b10a, points: 0, grade: false, solution: true -->
```python
def split_lyrics(lyrics: str, split_char: str) -> list[str]:
    ### BEGIN SOLUTION
    return lyrics.split(split_char)
    ### END SOLUTION
```

<!-- grade_id: b10t, points: 1, grade: true, solution: false -->
```python
assert split_lyrics("나는 사과, 당근, 바나나를 좋아해", ",") == ["나는 사과", " 당근", " 바나나를 좋아해"]
assert split_lyrics("반짝반짝 작은별 아름답게 비치네", " ") == ["반짝반짝", "작은별", "아름답게", "비치네"]
### BEGIN HIDDEN TESTS
assert split_lyrics("A|B|C|D", "|") == ["A", "B", "C", "D"]
### END HIDDEN TESTS
```