**주의사항**
 - 제출 전 커널을 재시작하고 **run all cells**를 눌러 모든 코드가 정상적으로 동작하는지 확인하세요.
 - 이름을 기입하고 Shift + Enter 를 누르세요.
 - 파일을 저장했는지 꼭 확인하세요.
 - 함수에서 `YOUR CODE HERE`라고 쓰여 있는 부분을 지우고 함수를 작성하세요.
 - 문제 아래에 테스트 코드를 실행해서 오류가 없는 지 확인해 보세요.
 - 테스트 코드를 통과했다고 해서 반드시 정답인 것은 아닙니다.

```python
이름 = ""
```

---

### 문제 1.

문자열에서 양쪽 끝에 있는 불필요한 공백이나 특정 문자를 제거하는 작업이 필요합니다. 이를 위해 `.strip()` 메서드를 사용하세요.

함수 `clean_string`를 작성하세요.

- 함수는 두 개의 매개변수를 받습니다: `original_str` (원본 문자열), `strip_chars` (제거할 문자. None이면 공백을 제거합니다)
- 함수는 `strip_chars`를 사용하여 `original_str`에서 양쪽 끝의 문자 또는 공백을 제거한 후, 그 결과 문자열을 반환합니다.

예를 들어:

- `original_str = "   안녕하세요   "`, `strip_chars = None`의 경우 출력은 `"안녕하세요"`이어야 합니다.
- `original_str = "!!!경고!!!"`, `strip_chars = "!"` 의 경우 출력은 ` "경고" `이어야 합니다.


<!-- grade_id: i1a, points: 0, grade: false, solution: true -->
```python
def clean_string(original_str, strip_chars=None):
    ### BEGIN SOLUTION
    if strip_chars:
        return original_str.strip(strip_chars)
    else:
        return original_str.strip()
    ### END SOLUTION
```

<!-- grade_id: i1t, points: 2, grade: true, solution: false -->
```python
assert clean_string("   안녕하세요   ", None) == "안녕하세요"
assert clean_string("!!!경고!!!", "!") == "경고"
### BEGIN HIDDEN TESTS
assert clean_string("   Hello World   ", None) == "Hello World"
### END HIDDEN TESTS
```

### 문제 2.

함수 `get_element`를 작성하세요.

주어진 리스트 `lst`와 인덱스 `index`를 사용하여 해당 위치의 요소를 반환하세요.

예를 들어:

- `lst = [10, 20, 30, 40, 50]`, `index = 2`의 경우 출력은 `30`이어야 합니다.

- `lst = ["안녕", "하세요", "파이썬"]`, `index = 0`의 경우 출력은 `"안녕"`이어야 합니다.


<!-- grade_id: i2a, points: 0, grade: false, solution: true -->
```python
def get_element(lst: list, index: int) -> any:
    ### BEGIN SOLUTION
    return lst[index]
    ### END SOLUTION
```

<!-- grade_id: i2t, points: 2, grade: true, solution: false -->
```python
assert get_element([10, 20, 30, 40, 50], 2) == 30
assert get_element(["안녕", "하세요", "파이썬"], 0) == "안녕"
assert get_element([True, False, True], 1) == False
assert get_element(["데이터", "베이스", "시스템"], 2) == "시스템"
### BEGIN HIDDEN TESTS
assert get_element(["이것", "은", "테스트", "입니다"], 3) == "입니다"
assert get_element([1, 2, 3, 4, 5, 6], 4) == 5
### END HIDDEN TESTS 
```

### 문제 3.

파이썬에서 튜플은 불변(immutable)입니다. 이는 튜플의 요소를 변경할 수 없다는 것을 의미합니다.

함수 `modify_tuple`를 작성하세요.

이 함수는 튜플 `tup`과 정수들 `index`와 `value`를 매개변수로 받습니다. 목표는 `tup`의 `index` 위치에 있는 값을 `value`로 바꾸는 것입니다. 그러나 튜플은 불변이므로 직접 변경할 수 없습니다. 따라서 함수는 새로운 튜플을 생성하여 해당 위치의 값을 변경하고 새 튜플을 반환해야 합니다.

예를 들어:

- `tup = (10, 20, 30, 40, 50)`, `index = 2`, `value = 100`의 경우 출력은 `(10, 20, 100, 40, 50)`이어야 합니다.

- `tup = ("안녕", "하세요", "파이썬")`, `index = 1`, `value = "다시"`의 경우 출력은 `("안녕", "다시", "파이썬")`이어야 합니다.


<!-- grade_id: i3a, points: 0, grade: false, solution: true -->
```python
def modify_tuple(tup: tuple, index: int, value: any) -> tuple:
    ### BEGIN SOLUTION
    lst = list(tup)
    lst[index] = value
    return tuple(lst)
    ### END SOLUTION
```

<!-- grade_id: i3t, points: 2, grade: true, solution: false -->
```python
assert modify_tuple((10, 20, 30, 40, 50), 2, 100) == (10, 20, 100, 40, 50)
assert modify_tuple(("안녕", "하세요", "파이썬"), 1, "다시") == ("안녕", "다시", "파이썬")
assert modify_tuple((True, False, True), 1, True) == (True, True, True)
assert modify_tuple(("데이터", "베이스", "시스템"), 2, "디자인") == ("데이터", "베이스", "디자인")
### BEGIN HIDDEN TESTS
assert modify_tuple(("이것", "은", "테스트"), 1, "정말") == ("이것", "정말", "테스트")
assert modify_tuple((1, 2, 3, 4, 5, 6, 7, 8, 9), 4, 0) == (1, 2, 3, 4, 0, 6, 7, 8, 9)
### END HIDDEN TESTS 
```

### 문제 4.

웹사이트에서 회원이 비밀번호를 잊어 버린 경우, 이를 다시 알려주는 기능을 만들고자 합니다. 아래에는 이메일 주소를 키로 하고 비밀번호를 값으로 하는 딕셔너리가 있습니다.

함수 `return_password`를 작성하세요.

- 함수는 두 개의 매개변수를 받습니다: `email_db` (이메일과 비밀번호를 저장하는 딕셔너리)와 `email` (확인하려는 이메일 주소).
- 함수는 이메일 주소가 `email_db`에 등록되어 있으면 비밀번호를, 그렇지 않으면 `False`를 반환합니다.

예를 들어:

- `email_db = {"alice@email.com": "1234", "bob@email.com": "abcd"}`, `email = "alice@email.com"`의 경우 출력은 `1234`이어야 합니다.
- `email_db = {"alice@email.com": "1234", "bob@email.com": "abcd"}`, `email = "charlie@email.com"`의 경우 출력은 `False`이어야 합니다.


<!-- grade_id: i4a, points: 0, grade: false, solution: true -->
```python
def return_password(email_db: dict, email: str) -> bool:
    ### BEGIN SOLUTION
    return email_db.get(email,False)
    ### END SOLUTION
```

<!-- grade_id: i4t, points: 2, grade: true, solution: false -->
```python
assert return_password({"alice@email.com": "1234", "bob@email.com": "abcd"}, "alice@email.com") == "1234"
assert return_password({"alice@email.com": "1234", "bob@email.com": "abcd"}, "charlie@email.com") == False
### BEGIN HIDDEN TESTS
assert return_password({"david@email.com": "qwerty", "emily@email.com": "asdfg"}, "emily@email.com") == "asdfg"
assert return_password({"david@email.com": "qwerty", "emily@email.com": "asdfg"}, "james@email.com") == False
### END HIDDEN TESTS
```

### 문제 5.

여러분은 편의점 사장님입니다. 편의점 관리를 위해 재고 관리하는 파이썬 프로그램을 만들고 있습니다. `update_inventory`라는 함수를 작성하세요.

- 함수는 두 개의 매개변수를 받습니다: `inventory` (제품명을 키로, 수량을 값으로 가지는 딕셔너리), `new_stock` (새로 들어온 제품과 수량을 가진 딕셔너리).
- 함수는 `new_stock` 딕셔너리의 모든 항목을 `inventory`에 추가합니다. 이미 존재하는 제품의 경우 수량을 업데이트합니다.
- 함수는 업데이트된 `inventory` 딕셔너리를 반환합니다.

예를 들어:

- `inventory = {"삼각김밥": 5, "우유": 10}`, `new_stock = {"삼각김밥": 3, "여명808": 7}`의 경우 출력은 `{"삼각김밥: 8, "우유": 10, "여명808": 7}`이어야 합니다.



<!-- grade_id: i5a, points: 0, grade: false, solution: true -->
```python
def update_inventory(inventory: dict, new_stock: dict) -> dict:
    ### BEGIN SOLUTION
    for item, quantity in new_stock.items():
        inventory[item] = inventory.get(item, 0) + quantity
    return inventory
    ### END SOLUTION
```

<!-- grade_id: i5t, points: 2, grade: true, solution: false -->
```python
assert update_inventory({"apple": 5, "banana": 10}, {"apple": 3, "cherry": 7}) == {"apple": 8, "banana": 10, "cherry": 7}
assert update_inventory({"milk": 2, "bread": 3}, {"bread": 1, "juice": 4}) == {"milk": 2, "bread": 4, "juice": 4}
### BEGIN HIDDEN TESTS
assert update_inventory({"coffee": 5, "tea": 7}, {"coffee": 2, "sugar": 1}) == {"coffee": 7, "tea": 7, "sugar": 1}
### END HIDDEN TESTS
```

### 문제 6.

당신은 웹사이트에서 회원가입 양식을 검증하는 코드를 작성하고 있습니다. 사용자로부터 입력 받은 정보는 딕셔너리에 저장됩니다. 검증해야 할 필드는 "username", "email", "age"입니다.

함수 `validate_form`를 작성하세요.

- 함수는 하나의 매개변수를 받습니다: `user_data` (사용자 정보를 담고 있는 딕셔너리)
- 함수는 딕셔너리 내용을 검증하여 새로운 딕셔너리를 생성합니다. 이때, 각 필드가 유효한지 여부 (True 또는 False)를 나타내어야 합니다. (예: 'username': True)
    - "username": 사용자 이름이 5글자 이상인지 검증
    - "email": '@' 문자가 포함되어 있는지 검증
    - "age": 나이가 18세 이상인지 검증

예를 들어:

- `user_data = {'username': 'john', 'email': 'john@email.com', 'age': 20}`의 경우 출력은 `{'username': False, 'email': True, 'age': True}`이어야 합니다.


<!-- grade_id: i6a, points: 0, grade: false, solution: true -->
```python
def validate_form(user_data: dict) -> dict:
    ### BEGIN SOLUTION
    return {k: len(v) >= 5 if k == 'username' else '@' in v if k == 'email' else v >= 18 for k, v in user_data.items()}
    ### END SOLUTION
```

<!-- grade_id: i6t, points: 2, grade: true, solution: false -->
```python
assert validate_form({'username': 'john', 'email': 'john@email.com', 'age': 20}) == {'username': False, 'email': True, 'age': True}
assert validate_form({'username': 'johnny', 'email': 'john.com', 'age': 16}) == {'username': True, 'email': False, 'age': False}
### BEGIN HIDDEN TESTS
assert validate_form({'username': 'michael', 'email': 'michael@email.com', 'age': 25}) == {'username': True, 'email': True, 'age': True}
### END HIDDEN TESTS
```

### 문제 7.

당신은 소셜 미디어 서비스의 개발자입니다. 각 사용자가 팔로우하는 사람들의 목록을 딕셔너리로 관리하려고 합니다.

함수 `create_following_dict`를 작성하세요.

- 함수는 두 개의 리스트를 매개변수로 받습니다: `users` (사용자 이름을 담은 리스트)와 `following` (팔로우하는 사람의 수를 담은 리스트).
- 두 리스트는 동일한 길이를 가집니다.
- 함수는 사용자 이름을 키로, 팔로우하는 사람의 수를 값으로 가지는 딕셔너리를 반환합니다.
- 단, 팔로우하는 사람이 5명 미만인 사용자는 딕셔너리에 포함되면 안 됩니다.

예를 들어:

- `users = ["Alice", "Bob", "Charlie"]`, `following = [3, 7, 5]`의 경우 출력은 `{'Bob': 7, 'Charlie': 5}`이어야 합니다.


<!-- grade_id: i7a, points: 0, grade: false, solution: true -->
```python
def create_following_dict(users: list, following: list) -> dict:
    ### BEGIN SOLUTION
    return {users[i]: following[i] for i in range(len(users)) if following[i] >= 5}
    ### END SOLUTION
```

<!-- grade_id: i7t, points: 2, grade: true, solution: false -->
```python
assert create_following_dict(["Alice", "Bob", "Charlie"], [3, 7, 5]) == {'Bob': 7, 'Charlie': 5}
assert create_following_dict(["David", "Emily", "Frank"], [2, 6, 1]) == {'Emily': 6}
### BEGIN HIDDEN TESTS
assert create_following_dict(["George", "Hannah"], [4, 8]) == {'Hannah': 8}
### END HIDDEN TESTS 
```

### 문제 8.

'무궁화 꽃이 피었습니다' 게임에서 참가자들은 무궁화 꽃이 피었을 때 움직이고, 꽃이 지면 정지해야 합니다. 참가자 중 무작위로 하나가 당첨자가 됩니다. 함수 `mugunghwa_game`를 작성하세요.

- 함수는 두 개의 매개변수를 받습니다:
    - n (게임을 진행할 횟수, 정수)
    - participants (게임에 참가하는 사람들의 이름을 담은 리스트)
- 함수는 무궁화 꽃이 피었을 때 1부터 n까지 숫자를 출력하고, n이 출력되고 나면 무작위로 선택된 당첨자의 이름을 출력합니다.
- 함수는 출력된 문자열 리스트를 반환합니다.
예를 들어:

`n = 3`, `participants = ['Alice', 'Bob', 'Charlie']`
출력은 `['1', '2', '3', 'Alice']` 또는 `['1', '2', '3', 'Bob']` 또는 `['1', '2', '3', 'Charlie']`

**주의**
1. 숫자가 아닌 문자열로 숫자가 출력된다는 것을 유의하세요.
2. choice 메소드를 잘 떠올려 보세요.


<!-- grade_id: i8a, points: 0, grade: false, solution: true -->
```python
import random

def mugunghwa_game(n: int, participants: list) -> list:
    ### BEGIN SOLUTION
    output = [str(i) for i in range(1, n + 1)]
    winner = random.choice(participants)
    output.append(winner)
    return output
    ### END SOLUTION
```

<!-- grade_id: i8t, points: 2, grade: true, solution: false -->
```python
assert mugunghwa_game(3, ['Alice', 'Bob', 'Charlie']) in [['1', '2', '3', 'Alice'], ['1', '2', '3', 'Bob'], ['1', '2', '3', 'Charlie']]
### BEGIN HIDDEN TESTS
assert mugunghwa_game(2, ['David', 'Emily']) in [['1', '2', 'David'], ['1', '2', 'Emily']]
### END HIDDEN TESTS
```

### 문제 9.

각 팀의 멤버 명단이 하나의 리스트로 주어지고, 모든 팀의 명단은 하나의 리스트에 묶여 있습니다. 모든 팀의 명단을 순서대로 출력하는 함수 `flatten_team_members`를 작성하세요.

- 함수는 하나의 매개변수를 받습니다: `project_teams` (모든 팀의 멤버 명단을 포함하는 2차원 리스트).
- 함수는 모든 팀의 멤버를 포함하는 하나의 '평평한' 리스트를 반환합니다. 리스트에 나타나는 이름의 순서는 원래의 2차원 리스트에 나타난 순서와 동일해야 합니다.

예를 들어,

- `project_teams = [["Ava", "Samantha", "James"], ["Lucille", "Zed"], ["Edgar", "Gabriel"]]`의 경우, 반환되는 리스트는 `["Ava", "Samantha", "James", "Lucille", "Zed", "Edgar", "Gabriel"]`이어야 합니다.


<!-- grade_id: i9a, points: 0, grade: false, solution: true -->
```python
def flatten_team_members(project_teams: list[list[str]]) -> list[str]:
    ### BEGIN SOLUTION
    flat_list = []
    for team in project_teams:
        for student in team:
            flat_list.append(student)
    return flat_list
    ### END SOLUTION
```

<!-- grade_id: i9t, points: 2, grade: true, solution: false -->
```python
assert flatten_team_members([["Ava", "Samantha", "James"], ["Lucille", "Zed"], ["Edgar", "Gabriel"]]) == ["Ava", "Samantha", "James", "Lucille", "Zed", "Edgar", "Gabriel"]
### BEGIN HIDDEN TESTS
assert flatten_team_members([["Alice"], ["Bob", "Charlie"], ["Dave"]]) == ["Alice", "Bob", "Charlie", "Dave"]
### END HIDDEN TESTS
```

### 문제 10.

`math` 모듈을 사용하여 두 점 사이의 거리를 계산하는 함수 `calculate_distance`를 작성하세요.

- 함수는 네 개의 매개변수를 받습니다: `x1, y1, x2, y2` (첫 번째 점의 x, y 좌표와 두 번째 점의 x, y 좌표).
- 함수는 두 점 사이의 거리를 반환합니다. 거리는 소수점 둘째 자리까지 반올림해야 합니다.

예를 들어:

- `x1 = 0, y1 = 0, x2 = 3, y2 = 4`의 경우 출력은 `5.0`이어야 합니다.

**힌트**: `math.sqrt` 함수와 `math.pow` 함수를 사용할 수 있습니다.



<!-- grade_id: i10a, points: 0, grade: false, solution: true -->
```python
import math
def calculate_distance(x1: float, y1: float, x2: float, y2: float) -> float:
    ### BEGIN SOLUTION
    distance = math.sqrt(math.pow(x2 - x1, 2) + math.pow(y2 - y1, 2))
    return round(distance, 2)
    ### END SOLUTION
```

<!-- grade_id: i10t, points: 2, grade: true, solution: false -->
```python
assert calculate_distance(0, 0, 3, 4) == 5.0
assert calculate_distance(1, 1, 4, 5) == 5.0
### BEGIN HIDDEN TESTS
assert calculate_distance(2, 2, 3, 3) == 1.41
### END HIDDEN TESTS
```